# MATH 666 Linear Algebra 1
## Assignment 1: Rank

* See ?? for a refresher on Python and Jupyter
* Run the tests to verify your solutions
* There are hidden tests that will be run once you submit your solutions. So your solutions may not be entirely correct even if they pass the tests below
* **Deadline:** Saturday, May 24 at 11:59 PM

```python
import numpy as np
```

Write code that calculates the rank of the matrix
$$ A = \begin{bmatrix} 1 & 2 & 1 \\ 3 & 0 & 5 \\ 7 & -1 & 0 \end{bmatrix}. $$
Save your answers as <code>A</code> and <code>rank_A</code>, respectively.

<!-- grade_id: cell-7797be35b6115b60, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[2], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-5c28a8a9c607f211, points: 1, grade: true, solution: false -->
```python
"Verify that A and rank_A have been computed correctly. This cell contains hidden tests. (1 mark)"
assert rank_A.is_integer() and rank_A>=0, "rank_A should be a nonnegative integer."
print("Problem 1: Success!")
```

Write code that calculates the rank of the matrix
$$ [A\,|\,\textbf{b}] = \left[\begin{array}{ccc|c} 1 & 2 & 1 & 0 \\ 3 & 0 & 5 & -1 \\ 7 & -1 & 0 & 2\end{array}\right].$$
Save your answers as <code>B</code> and <code>rank_B</code>, respectively.

<!-- grade_id: cell-4ed90714b7792558, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c813104679190268, points: 1, grade: true, solution: false -->
```python
"Verify that B and rank_B have been computed correctly. This cell contains hidden tests. (1 mark)"
assert rank_B.is_integer() and rank_B>=0, "rank_B should be a nonnegative integer."
print("Problem 2: Success!")
```

How many solutions does the following system of equations have?
$$
\begin{align*}
x+2y+z &= 0 \\
3x+5z &= -1 \\
7x-y &= 2
\end{align*}
$$
If there are none, enter <code>S=0</code>, if there is exactly one, enter <code>S=1</code>, and if there are infinitely many, enter <code>S=np.inf</code>

<!-- grade_id: cell-a124372e413cfcc6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ae67f2bf581e451e, points: 1, grade: true, solution: false -->
```python
"This cell contains hidden tests. (1 mark)"
print("Problem 3:Success!")
```